Alohan'ny mamerina dia avereno atao Run ny notebook iray manontolo. Ny fanaovana azy dia redémarrena mihitsy ny kernel aloha (jereo menubar, safidio **Kernel$\rightarrow$Restart Kernel and Run All Cells**).

Izay misy hoe `YOUR CODE HERE` na "YOUR ANSWER HERE" ihany no fenoina. Afaka manampy cells vaovao raha ilaina. Aza adino ny mameno references eo ambany raha ilaina.

## References
Eto ilay references rehetra no apetraka

---

```python
from random import randrange
import numpy as np
from sklearn.metrics import mean_squared_error, log_loss
from sklearn.datasets import load_breast_cancer, load_diabetes


def grad_check_sparse(f, x, analytic_grad, num_checks=10, h=1e-5, error=1e-9):
    """
    sample a few random elements and only return numerical
    in this dimensions
    """

    for i in range(num_checks):
        ix = tuple([randrange(m) for m in x.shape])

        oldval = x[ix]
        x[ix] = oldval + h  # increment by h
        fxph = f(x)  # evaluate f(x + h)
        x[ix] = oldval - h  # increment by h
        fxmh = f(x)  # evaluate f(x - h)
        x[ix] = oldval  # reset

        grad_numerical = (fxph - fxmh) / (2 * h)
        grad_analytic = analytic_grad[ix]
        rel_error = abs(grad_numerical - grad_analytic) / (
            abs(grad_numerical) + abs(grad_analytic)
        )
        print(
            "numerical: %f analytic: %f, relative error: %e"
            % (grad_numerical, grad_analytic, rel_error)
        )
        assert rel_error < error

def rel_error(x, y):
    """ returns relative error """
    return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))
```

# Linear regression

```python
data = load_diabetes()
X_train1, y_train1 = data.data, data.target
w1 = np.random.randn(X_train1.shape[1]) * 0.0001
b1 = np.random.randn(1) * 0.0001
```

<!-- grade_id: cell-5ffdf99ad7cdfd69, points: 0, grade: false, solution: true -->
```python
def mse_loss_naive(w, b, X, y, alpha=0):
    """
    MSE loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    # YOUR CODE HERE
    for i in range(len(X)):
        loss += (y[i] - (w.transpose()).dot(X[i]) -  b) ** 2
        dw -= (y[i] - (w.transpose()).dot(X[i]) -  b) * X[i]
        db -= y[i] - (w.transpose()).dot(X[i]) -  b
    loss = loss / X.shape[0]
    dw = 2*(dw / X.shape[0])
    db = 2*(db/ X .shape[0])
    
    return loss, dw, np.array(db).reshape(1,)
```

## Naive Linear regression loss

<!-- grade_id: cell-079e42b153b67a60, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  3.1281544173592984e-16
Gradient check w
numerical: -3.234110 analytic: -3.234110, relative error: 4.608004e-08
numerical: -1.553189 analytic: -1.553189, relative error: 1.796933e-09
numerical: -3.153319 analytic: -3.153319, relative error: 8.641310e-09
numerical: 2.892059 analytic: 2.892061, relative error: 3.783147e-07
numerical: -1.376396 analytic: -1.376395, relative error: 3.055484e-07
numerical: -3.153319 analytic: -3.153319, relative error: 8.641310e-09
numerical: -1.275045 analytic: -1.275046, relative error: 1.263567e-07
numerical: -0.315455 analytic: -0.315455, relative error: 8.276968e-07
numerical: -1.376396 analytic: -1.376395, relative error: 3.055484e-07
numerical: -1.553189 analytic: -1.553189, relative error: 1.796933e-09
numerical: -4.145420 analytic: -4.145419, relative error: 1.051224e-07
numerical: -4.296085 analytic: -4.296088, relative error: 3.792221e-07
numerical: 2.892059 analytic: 2.892061, relative error: 3.783147e-07
numerical: -1.275045 analyt
```

## Naive Ridge regression loss

<!-- grade_id: cell-14a64cb1cfe70b98, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -3.234110 analytic: -3.234110, relative error: 4.608004e-08
numerical: -1.275045 analytic: -1.275046, relative error: 1.263567e-07
numerical: -1.376396 analytic: -1.376395, relative error: 3.055484e-07
numerical: -1.275045 analytic: -1.275046, relative error: 1.263567e-07
numerical: -4.145420 analytic: -4.145419, relative error: 1.051224e-07
numerical: -1.275045 analytic: -1.275046, relative error: 1.263567e-07
numerical: -0.315455 analytic: -0.315455, relative error: 8.276968e-07
numerical: -2.801913 analytic: -2.801914, relative error: 1.522192e-07
numerical: -4.145420 analytic: -4.145419, relative error: 1.051224e-07
numerical: -0.315455 analytic: -0.315455, relative error: 8.276968e-07
numerical: -1.553189 analytic: -1.553189, relative error: 1.796933e-09
numerical: -3.234110 analytic: -3.234110, relative error: 4.608004e-08
numerical: -1.376396 analytic: -1.376395, relative error: 3.055484e-07
numerical: -2.801913 analytic: -2.801914, relative error: 1.
```

<!-- grade_id: cell-1528a28f467d90c7, points: 0, grade: false, solution: true -->
```python
def mse_loss_vectorized(w, b, X, y, alpha=0):
    """
    MSE loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    # YOUR CODE HERE
    y_pred = X.dot(w) + b
    
    # loss
    loss = (1/y.shape[0]) * ((y_pred - y).dot(y_pred - y) + alpha * ((w.transpose()).dot(w)))
    
    # gradient with respect to weights w
    dw = (2/X.shape[0]) * ((alpha * w) + (X.transpose()).dot(y_pred - y))
    
    # gradient with respect to bias
    db = (2/X.shape[0]) * np.sum(y_pred - y)
    
    return loss, dw, np.array(db).reshape(1,)
```

## Vectorised Linear regression loss

<!-- grade_id: cell-41637ca21c8f938d, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  2.502523533887439e-16
Gradient check w
numerical: 2.892061 analytic: 2.892061, relative error: 9.389016e-10
numerical: -3.234111 analytic: -3.234110, relative error: 9.452968e-08
numerical: -2.801914 analytic: -2.801914, relative error: 1.007962e-08
numerical: -1.275046 analytic: -1.275046, relative error: 1.630403e-08
numerical: -3.234111 analytic: -3.234110, relative error: 9.452968e-08
numerical: -1.376395 analytic: -1.376395, relative error: 2.484166e-08
numerical: -1.376395 analytic: -1.376395, relative error: 2.484166e-08
numerical: -4.145419 analytic: -4.145419, relative error: 4.845586e-08
numerical: -1.376395 analytic: -1.376395, relative error: 2.484166e-08
numerical: -3.153319 analytic: -3.153319, relative error: 8.641310e-09
numerical: -3.153319 analytic: -3.153319, relative error: 8.641310e-09
numerical: -3.234111 analytic: -3.234110, relative error: 9.452968e-08
numerical: -2.801914 analytic: -2.801914, relative error: 1.007962e-08
numerical: -4.296088 analy
```

## Vectorized ridge regression loss

<!-- grade_id: cell-bb0ff99c19a85e0f, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -1.553190 analytic: -1.553190, relative error: 2.689604e-08
numerical: -1.376395 analytic: -1.376396, relative error: 3.644297e-08
numerical: -4.145419 analytic: -4.145419, relative error: 8.666912e-08
numerical: -2.801913 analytic: -2.801914, relative error: 2.921326e-08
numerical: -2.801913 analytic: -2.801914, relative error: 2.921326e-08
numerical: -3.153319 analytic: -3.153319, relative error: 3.369887e-08
numerical: -4.296088 analytic: -4.296088, relative error: 6.194543e-08
numerical: -1.275047 analytic: -1.275047, relative error: 2.265001e-08
numerical: -4.296088 analytic: -4.296088, relative error: 6.194543e-08
numerical: -4.145419 analytic: -4.145419, relative error: 8.666912e-08
numerical: 2.892062 analytic: 2.892062, relative error: 4.981050e-09
numerical: -4.145419 analytic: -4.145419, relative error: 8.666912e-08
numerical: -2.801913 analytic: -2.801914, relative error: 2.921326e-08
numerical: -3.234110 analytic: -3.234110, relative error: 1.03
```

# Logistic regression

```python
def sigmoid(z):
    return 1 / (1 + np.exp(-z))

data = load_breast_cancer()
X_train2, y_train2 = data.data, data.target
w2 = np.random.randn(X_train2.shape[1]) * 0.0001
b2 = np.random.randn(1) * 0.0001
```

# Naive

<!-- grade_id: cell-6ec156568b0c6e29, points: 0, grade: false, solution: true -->
```python
def log_loss_naive(w, b, X, y, alpha=0):
    """
    log loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    # YOUR CODE HERE
    
    def logistic(y, m):
        res = y * np.log(m) + (1-y) * np.log(1-m)
        return res
    
    for i in range(X.shape[0]):
        sig = sigmoid(w.T.dot(X[i]) + b)
        loss -= logistic(y[i], sig)
    
    loss += alpha * (w.transpose()).dot(w)
    loss = (1/X.shape[0]) * loss
    
    for i in range(w.shape[0]):
        x1 = 2 * alpha * w[i]
        for j in range(X.shape[0]):
            x1 += (sigmoid((w.transpose()).dot(X[j]) + b) - y[j]) * X[j][i]
        dw[i] = x1 / y.shape[0]
    
    x1 = 0
    for i in range(X.shape[0]):
        x1 += sigmoid(((w.transpose()).dot(X[i])) + b) - y[i]
    db = (1/X.shape[0]) * x1
    
    
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-d3c078eb2449ee61, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  7.951603207737165e-17
Gradient check w
numerical: -0.138444 analytic: -0.138444, relative error: 6.360337e-11
numerical: 7.877575 analytic: 7.877576, relative error: 2.331424e-08
numerical: -0.006550 analytic: -0.006550, relative error: 2.615908e-09
numerical: 0.003799 analytic: 0.003799, relative error: 5.990420e-10
numerical: 0.012817 analytic: 0.012817, relative error: 2.341918e-09
numerical: -0.019883 analytic: -0.019883, relative error: 2.517171e-10
numerical: 0.003799 analytic: 0.003799, relative error: 5.990420e-10
numerical: 0.009376 analytic: 0.009376, relative error: 3.169600e-10
numerical: -0.000324 analytic: -0.000324, relative error: 1.030730e-08
numerical: -0.010112 analytic: -0.010112, relative error: 2.130841e-10
numerical: -0.006550 analytic: -0.006550, relative error: 2.615908e-09
numerical: 50.452457 analytic: 50.452618, relative error: 1.599717e-06
numerical: 0.000194 analytic: 0.000194, relative error: 1.201145e-07
numerical: 0.032599 analytic: 0.0325
```

# Naive with regulariztion

<!-- grade_id: cell-b91805308ad91ec5, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: 50.452457 analytic: 50.452618, relative error: 1.599717e-06
numerical: -0.019884 analytic: -0.019884, relative error: 4.009081e-10
numerical: -1.348644 analytic: -1.348644, relative error: 6.944073e-08
numerical: -0.000081 analytic: -0.000081, relative error: 2.057767e-08
numerical: -0.305242 analytic: -0.305242, relative error: 1.022863e-09
numerical: -0.010112 analytic: -0.010112, relative error: 5.582703e-10
numerical: 50.452457 analytic: 50.452618, relative error: 1.599717e-06
numerical: 50.452457 analytic: 50.452618, relative error: 1.599717e-06
numerical: -0.008306 analytic: -0.008306, relative error: 8.055911e-10
numerical: 0.236599 analytic: 0.236599, relative error: 6.657039e-11
numerical: 0.236599 analytic: 0.236599, relative error: 6.657039e-11
numerical: -0.000081 analytic: -0.000081, relative error: 2.057767e-08
numerical: -0.015824 analytic: -0.015824, relative error: 1.601730e-11
numerical: -0.006550 analytic: -0.006550, relative error: 2.6748
```

# Vectorized

<!-- grade_id: cell-a96e9a6d51919ffc, points: 0, grade: false, solution: true -->
```python
def log_loss_vectorized(w, b,X, y, alpha=0):
    """
    log loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    """
    loss = 0.0
    dw = np.zeros_like(w)
    
    # YOUR CODE HERE
    sig = sigmoid(X.dot(w) + b)
    
    loss -= sum(y * np.log(sig) + (1 - y) * np.log(1 - sig))
    loss = (loss + alpha * w.transpose().dot(w)) / X.shape[0]
    
    dw = ((2 * alpha * w) + (X.transpose().dot(sig - y))) / X.shape[0]
    
    db = np.sum(sig - y) / X.shape[0]
    
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-ca14e49ada130789, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  1.590320641547433e-16
Gradient check w
numerical: -1.496190 analytic: -1.496190, relative error: 9.438344e-10
numerical: 0.032599 analytic: 0.032599, relative error: 3.399354e-11
numerical: 7.877575 analytic: 7.877576, relative error: 2.331459e-08
numerical: -0.000324 analytic: -0.000324, relative error: 1.753830e-09
numerical: -0.000277 analytic: -0.000277, relative error: 1.498908e-08
numerical: 0.009376 analytic: 0.009376, relative error: 3.169600e-10
numerical: -0.138444 analytic: -0.138444, relative error: 6.360337e-11
numerical: -1.280233 analytic: -1.280233, relative error: 4.659127e-10
numerical: 0.017385 analytic: 0.017385, relative error: 1.116616e-09
numerical: 7.877575 analytic: 7.877576, relative error: 2.331459e-08
numerical: -0.305242 analytic: -0.305242, relative error: 1.025588e-09
numerical: -0.138444 analytic: -0.138444, relative error: 6.360337e-11
numerical: 0.012817 analytic: 0.012817, relative error: 2.341919e-09
numerical: -0.000277 analytic: -0.00
```

# Vectorized with regularization

<!-- grade_id: cell-bce082fccf8b7057, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: 0.032599 analytic: 0.032599, relative error: 3.671414e-11
numerical: 108.359010 analytic: 108.359550, relative error: 2.488701e-06
numerical: -0.000887 analytic: -0.000887, relative error: 3.477619e-08
numerical: 0.236599 analytic: 0.236599, relative error: 6.657051e-11
numerical: 0.236599 analytic: 0.236599, relative error: 6.657051e-11
numerical: 0.003799 analytic: 0.003799, relative error: 3.317074e-10
numerical: 0.017385 analytic: 0.017385, relative error: 1.285784e-09
numerical: -0.019884 analytic: -0.019884, relative error: 2.613203e-10
numerical: -0.019884 analytic: -0.019884, relative error: 2.613203e-10
numerical: -0.015824 analytic: -0.015824, relative error: 1.601730e-11
numerical: -1.280233 analytic: -1.280233, relative error: 4.649856e-10
numerical: 0.017069 analytic: 0.017069, relative error: 1.242419e-09
numerical: -0.015824 analytic: -0.015824, relative error: 1.601730e-11
numerical: -0.006550 analytic: -0.006550, relative error: 3.098654e-09
```

# Gradient descent for Linear models

<!-- grade_id: cell-485e52c0efd4f4a9, points: 0, grade: false, solution: true -->
```python
class LinearModel():
    def __init__(self):
        self.w = None
        self.b = None

    def train(self, X, y, learning_rate=1e-3, alpha=0, num_iters=100, batch_size=200, verbose=False):
        N, d = X.shape
        
        if self.w is None: # Initialization
            self.w = 0.001 * np.random.randn(d)
            self.b = 0.0

        # Run stochastic gradient descent to optimize w
        
        loss_history = []
        for it in range(num_iters):
            X_batch = None
            y_batch = None
                                                               
            # Sample batch_size elements in X_batch and y_batch
            # X_batch shape is  (batch_size, d) and y_batch shape is (batch_size,)                                                                                          
            # Hint: Use np.random.choice to generate indices
            # YOUR CODE HERE
            rand = np.random.choice(N, batch_size, replace=False)
            
            X_batch = X[rand,:]
            y_batch = y[rand]
            
            # evaluate loss and gradient
            loss, dw, db = self.loss(X_batch, y_batch, alpha)
            loss_history.append(loss)

            # perform parameter update                                                                
            # Update the weights w and bias b using the gradient and the learning rate.          
            # YOUR CODE HERE
            self.w -= learning_rate * dw
            self.b -= learning_rate * db
            
            if verbose and it % 10000 == 0:
                print("iteration %d / %d: loss %f" % (it, num_iters, loss))
                
        return loss_history

    def predict(self, X):
        pass

    def loss(self, X_batch, y_batch, reg):
        pass

class LinearRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return mse_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        # YOUR CODE HERE
        return X.dot(self.w) + self.b

class LogisticRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return log_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        """ Return prediction labels vector of 0 or 1 """
        # YOUR CODE HERE
        y = X.dot(self.w) + self.b
        r = np.zeros_like(y)
        for i in range(len(y)):
            r[i] = 1 if y[i] > 0 else 0
        return r
```

## Linear regression with gradient descent

<!-- grade_id: cell-92f36a3b387a4277, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LinearRegression

sk_model = LinearRegression(fit_intercept=True)
sk_model.fit(X_train1, y_train1)
sk_pred = sk_model.predict(X_train1)
sk_mse = mean_squared_error(sk_pred, y_train1)

model = LinearRegressor()
model.train(X_train1, y_train1, num_iters=75000, batch_size=64, learning_rate=1e-2, verbose=True)
pred = model.predict(X_train1)
mse = mean_squared_error(pred, y_train1)

print("MSE scikit-learn:", sk_mse)
print("MSE gradient descent model :", mse)
assert mse - sk_mse < 100
```

Output:
```
iteration 0 / 75000: loss 27038.202583
iteration 10000 / 75000: loss 3693.466885
iteration 20000 / 75000: loss 2446.433191
iteration 30000 / 75000: loss 2670.691459
iteration 40000 / 75000: loss 2910.608874
iteration 50000 / 75000: loss 2821.783531
iteration 60000 / 75000: loss 2770.712888
iteration 70000 / 75000: loss 2425.398640
MSE scikit-learn: 2859.6963475867506
MSE gradient descent model : 2885.6788695293612
```

## Logistc regression with gradient descent

<!-- grade_id: cell-925a2ddb5c1ba7ff, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LogisticRegression
from sklearn.preprocessing import StandardScaler

scaler = StandardScaler()
X_train2 = scaler.fit_transform(X_train2)

sk_model = LogisticRegression(fit_intercept=True)
sk_model.fit(X_train2, y_train2)
sk_pred = sk_model.predict(X_train2)
sk_log_loss = log_loss(sk_pred, y_train2)

model = LogisticRegressor()
model.train(X_train2, y_train2, num_iters=75000, batch_size=64, learning_rate=1e-3, verbose=True)
pred = model.predict(X_train2)
model_log_loss = log_loss(pred, y_train2)

print("Log-loss scikit-learn:", sk_log_loss)
print("Log-loss gradiet descent model :", model_log_loss)
print("Error :", rel_error(sk_log_loss, model_log_loss))
assert rel_error(sk_log_loss, model_log_loss) < 1e-7
```

Output:
```
iteration 0 / 75000: loss 0.692649
iteration 10000 / 75000: loss 0.077008
iteration 20000 / 75000: loss 0.056454
iteration 30000 / 75000: loss 0.079893
iteration 40000 / 75000: loss 0.096626
iteration 50000 / 75000: loss 0.030329
iteration 60000 / 75000: loss 0.071384
iteration 70000 / 75000: loss 0.019341
Log-loss scikit-learn: 0.4249086712816093
Log-loss gradiet descent model : 0.4249086712816093
Error : 0.0
```

```python

```