Alohan'ny mamerina dia avereno atao Run ny notebook iray manontolo. Ny fanaovana azy dia redémarrena mihitsy ny kernel aloha (jereo menubar, safidio **Kernel$\rightarrow$Restart Kernel and Run All Cells**).

Izay misy hoe `YOUR CODE HERE` na "YOUR ANSWER HERE" ihany no fenoina. Afaka manampy cells vaovao raha ilaina. Aza adino ny mameno references eo ambany raha ilaina.

## References
Eto ilay references rehetra no apetraka

# https://www.kaggle.com/code/kunjmehta/gaussian-kernel-regression-from-scratch
# https://github.com/annonymousecode/Model-Zoology-and-Neural-Taskonomy

```python
import cvxpy as cp
import pandas as pd
import numpy as np
from sklearn.datasets import load_diabetes
from sklearn.kernel_ridge import KernelRidge
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

def rel_error(x, y):
    """ returns relative error """
    return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))

data = load_diabetes()
X_train, y_train = data.data, data.target
```

# Linear Program

The minimum amount of nutrients of a healthy diet is presented in the following table:
|  Nutrient | Minimum recommended intake |  
|---|---|
|  Calories | 3 kcal  |  
|  Protein | 70 g  |   
|  Calcium | 0.8 g  |  
|  Iron | 12 mg  | 
|  Vitamin A | 5000 IU  | 
|  Vitamin B1 | 1.8 mg  | 
|  Vitamin B2 | 2.7 mg  | 
|  Niacin | 18 mg  | 
|  Vitamin C | 75 mg  | 

Each row of the table in the *diet.csv* file presents the amount of nutrients present in one dollar worth of food. How many dollars should one spend on each food so that the total cost is minimum ?

<!-- grade_id: cell-36f73ab369548dc2, points: 0, grade: false, solution: true -->
```python
df = pd.read_csv("diet.csv", index_col=0)
opt_solution = None
# YOUR CODE HERE
donnees = df.to_numpy()

X = cp.Variable(donnees.shape[0])
b = np.array([3000,70,0.8,12,5000,1.8,2.7,18,75])

prob = cp.Problem(cp.Minimize(cp.sum(donnees.T @ X)), [donnees.T @ X >= b, X >=0 ])
result = prob.solve()

opt_solution = X.value
print(opt_solution)
```

Output:
```
[-1.52910474e-08 -1.87604251e-08 -2.05207909e-08  9.16104169e-09
 -1.55353042e-08 -1.72436916e-08 -1.86451206e-08 -1.43156216e-08
 -1.79285743e-08 -1.95732502e-08 -1.51635265e-08  1.46339494e-08
 -1.52879137e-08 -1.10829417e-08 -1.51686892e-08  1.12556522e+02
  4.90373154e-06 -7.90067633e-09 -1.62071084e-08  1.62556125e-07
 -2.07795338e-08  4.77846935e-07  4.27382417e+01  2.14642307e-06
 -1.79979309e-08 -1.97669601e-08 -1.01779440e-08 -2.43053939e-08
 -1.67307443e-08 -2.12942132e-08 -1.78782040e-08 -9.17550697e-09
  2.14933768e-07  5.19986721e-08 -1.65433094e-08  1.74878500e-08
  4.03751684e-09 -2.06007270e-08 -1.47741420e-08 -2.09350453e-08
  2.08543771e-07  5.86222892e-08  1.08239948e-07  8.84602103e-08
  1.52445902e-08  1.74150652e-08  7.10739915e-08  7.86180218e-08
  5.56265890e-07  1.56460024e-08  8.15477117e-09  2.72227508e-02
  4.29596704e-08  1.23596175e-07  6.32796060e-07  3.42101734e-07
  1.45632883e-07  3.35189560e-08 -1.90364159e-08 -1.68454817e-09
  1.95509809e-08  4.92387
```

```python
df
```

Output:
```
Calories (1000)  Protein (g)  Calcium (g)  Iron (mg)  \
Wheat Flour (Enriched)              44.7         1411          2.0        365   
Macaroni                            11.6          418          0.7         54   
Wheat Cereal (Enriched)             11.8          377         14.4        175   
Corn Flakes                         11.4          252          0.1         56   
Corn Meal                           36.0          897          1.7         99   
...                                  ...          ...          ...        ...   
Chocolate                            8.0           77          1.3         39   
Sugar                               34.9            0          0.0          0   
Corn Syrup                          14.7            0          0.5         74   
Molasses                             9.0            0         10.3        244   
Strawberry Preserves                 6.4           11          0.4          7   

                         Vitamin A (1000 IU)  Thiami
```

<!-- grade_id: cell-efdad2e3d376e010, points: 4, grade: true, solution: false -->
```python

```

# Kernel Ridge Regression

<!-- grade_id: cell-210d2ae0b70955f7, points: 0, grade: false, solution: true -->
```python
def gaussian_kernel(x1, x2, sigma=1.0):
    # YOUR CODE HERE
    s = np.exp(- np.linalg.norm(x1 - x2)**2 / 2*sigma**2)
    return s

class KernelRidgeRegression():
    def __init__(self, kernel, lambd=1.0):
        self.lambd = lambd # regularization coefficient
        self.alpha = None # dual variable
        self.kernel = kernel # kernel function
    
    def fit(self, X, y):
        # YOUR CODE HERE
        N = X.shape[0]
        I = np.eye(N)
        
        #Kernel function
        k = np.zeros((N,N))
        for i in range(N):
            for j in range(N):
                k[i][j] = self.kernel(X[i,:],X[j,:])
        
        temp = np.linalg.inv(k+self.lambd*I)
        self.alpha = y@temp     
    
    def predict(self, X):
        # YOUR CODE HERE
        N = X.shape[0]
        y_pred = np.zeros(N)
        for i in range(N):
            k = np.zeros(N)
            for j in range(N):
                k[j] = self.kernel(X[i,:],X[j,:])
            y_pred[i] = self.alpha@k
        return y_pred 
```

<!-- grade_id: cell-7d655e9354dd8a5a, points: 4, grade: true, solution: false -->
```python
kr = KernelRidgeRegression(kernel=gaussian_kernel)
kr.fit(X_train,y_train)
y_pred = kr.predict(X_train)

sklearn_kr = KernelRidge(kernel=gaussian_kernel)
sklearn_kr.fit(X_train,y_train)
sklearn_y_pred = kr.predict(X_train)

error = rel_error(kr.alpha, sklearn_kr.dual_coef_)
print("error: ", error)
assert error <= 1e-11

error_pred = rel_error(y_pred, sklearn_y_pred)
print("prediction error: ", error_pred)
assert error_pred <= 1e-11
```

Output:
```
error:  4.918238903386735e-13
prediction error:  0.0
```

# Non-negative Linear Regression

<!-- grade_id: cell-37e7efb4bf96fe1a, points: 0, grade: false, solution: true -->
```python
class NonNegativeLinearRegression:
    def __init__(self, fit_intercept=True):
        self.w = None
        self.fit_intercept = fit_intercept
    
    def fit(self, X, y):
        # YOUR CODE HERE
        if self.fit_intercept:
            X = np.concatenate((np.ones((X.shape[0],1)), X),axis=1)
        x = cp.Variable(X.shape[1])
        problem = cp.sum_squares(X @ x - y)
        prob = cp.Problem(cp.Minimize(problem), [x >= 0])
        prob.solve()
        self.w = x.value
    
    def predict(self, X):
        # YOUR CODE HERE
        if self.fit_intercept == True:
            X = np.concatenate((np.ones((X.shape[0],1)), X),axis=1)
        return X.dot(self.w)
```

<!-- grade_id: cell-b1a1983846be3759, points: 4, grade: true, solution: false -->
```python
nnlr = NonNegativeLinearRegression(fit_intercept=True)
nnlr.fit(X_train,y_train)
y_pred = kr.predict(X_train)
mse = mean_squared_error(y_pred, y_train)

sk_model = LinearRegression(fit_intercept=True, positive=True)
sk_model.fit(X_train, y_train)
sk_pred = sk_model.predict(X_train)
sk_mse = mean_squared_error(sk_pred, y_train)

print("MSE scikit-learn:", sk_mse)
print("MSE gradient descent model :", mse)
assert mse - sk_mse < 200
```

Output:
```
MSE scikit-learn: 3074.1786797315144
MSE gradient descent model : 3270.2512240824185
```

```python

```